---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Getting started with Thresholding and Histogram Equalization

```yaml
type: VideoExercise
key: 1129c6f943
xp: 50
```

`@projector_key`
812b4bba43f5ff9f855477093861f158
